Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Evan Huynh"
COLLABORATORS = ""
```

---

## Lab: basic Python syntax


#### Exercise 1
Write a function 'mymax' that accepts two numbers as parameters and that prints the larger number to the screen. For this elementary exercise, do not use the built-in function 'max()' - just revert to elementary expressions and statements.

<!-- grade_id: cell-ee8c56c7cc588b03, points: 0, grade: false, solution: true -->
```python
def mymax(x, y):
    """
    Returns the maximum of x and y
    """
    # YOUR CODE HERE
    return x if x > y else y
    ################


print(mymax(2, 6))
print(mymax(5, 3))
```

Output:
```
6
5
```

<!-- grade_id: cell-304328f9c7d4c051, points: 1, grade: true, solution: false -->
```python
assert mymax(2, 6) == 6, f'mymax(2, 6) -> {mymax(2, 6)} != 6'
assert mymax(5, 3) == 5, f'mymax(5, 3) -> {mymax(5, 3)} != 5'
```

<!-- grade_id: cell-3de8d447b6562bf0, points: 1, grade: true, solution: false -->
```python
import unittest.mock
with unittest.mock.patch('__main__.max') as mockmax:
    mymax(2, 6)
mockmax.assert_not_called()
```

#### Exercise 2

Write a function to sum the elements of a sequence with a for loop. You can assume that it consists of numeric values. Do not use the builtin sum function

<!-- grade_id: cell-5ac58bb5e671db91, points: 0, grade: false, solution: true -->
```python
def mysum(sequence):
    """
    Returns the sum of the elements in sequence

    I have had some fun with it using lambda expression.
    Source: https://stackoverflow.com/questions/20253651/recursively-sum-the-integers-in-an-array
    """
    # YOUR CODE HERE
    return sequence[0] if len(sequence) == 1 else sequence[0] + mysum(sequence[1::])
    ################


print(mysum([1, 2]))
print(mysum([1, 1/2, 1/4]))
print(mysum([2, 6]))
```

Output:
```
3
1.75
8
```

<!-- grade_id: cell-24ac435261756312, points: 1, grade: true, solution: false -->
```python
assert mysum([1, 2]) == 3, f'mysum([1, 2]) -> {mysum([1, 2])} != 3'
assert mysum((1, 1/2, 1/4)) == 1.75, f'mysum((1, 1/2, 1/4)) -> {mysum((1, 1/2, 1/4))} != 1.75'
```

<!-- grade_id: cell-f3e385c711f70755, points: 1, grade: true, solution: false -->
```python
import unittest.mock
with unittest.mock.patch('__main__.sum') as mocksum:
    mysum((2, 6))
mocksum.assert_not_called()
```

#### Exercise 3

Variables of `str` type has a method `join` defined such that when a sequence (e.g. `list`, `tuple`) of strings is given as an argument, they are concatenated to a single string with the variable as a separator. E.g.

    >>> separator = '-'
    >>> separator.join(['father', 'in', 'law'])
    'father-in-law'

Use the `join` to generate a histogram with stars such that 

    histogram([3, 4, 1]) 
    
returns

    ***
    ****
    *



<!-- grade_id: cell-631ae09a389a66d4, points: 0, grade: false, solution: true -->
```python
def histogram(sequence):
    """
    Returns histogram in the form of characters
    
    >>> histogram([3, 4, 1])
    ***
    ****
    *
    
    """
    # YOUR CODE HERE
    return "\n".join(["*" * _ for _ in sequence])
    ################
    

print(f'Example: the histogram of [3, 4, 1] is \n{histogram([3, 4, 1])}')
```

Output:
```
Example: the histogram of [3, 4, 1] is 
***
****
*
```

<!-- grade_id: cell-901ea19debe3ee4c, points: 1, grade: true, solution: false -->
```python
assert histogram([]) == "", f'"{histogram([])}" != "" '
assert histogram([1, 2]) == "*\n**", f'"{histogram([1, 2])}" != \n"*\n**"'
assert histogram([1, 2, 3]) == "*\n**\n***", f'"{histogram([1, 2, 3])}" != \n "*\n**\n***"'
```

#### Exercise 4

Compute how many times a character is present in a string. Use a dictionary to represent the string. 

<!-- grade_id: cell-94b1ecc182583f9b, points: 0, grade: false, solution: true -->
```python
def character_count(text):
    """
    Returns character count of text
    Input: text of type str
    Output: dict that maps a character to a number count

    Source:
        https://stackoverflow.com/questions/5664494/letter-count-dictionary
        https://stackoverflow.com/questions/40950905/find-count-of-characters-within-the-string-in-python
    """
    # YOUR CODE HERE
    return {e : text.count(e) for e in text}
    ################


print(f"Example: Character count of 'ABBA': {character_count('ABBA')} \nExpected {{'A': 2, 'B': 2}}")
```

Output:
```
Example: Character count of 'ABBA': {'A': 2, 'B': 2} 
Expected {'A': 2, 'B': 2}
```

<!-- grade_id: cell-2a787760db753909, points: 1, grade: true, solution: false -->
```python
calculated = character_count("ABBA")
expected = {"A": 2, "B": 2}
assert calculated == expected, f'{calculated} != {expected}'

calculated = character_count("CCGGAAGAGCTTACTTAG")
expected = {"A": 5, "C": 4, "G": 5, "T": 4}
assert calculated == expected, f'{calculated} != {expected}'
```

#### Exercise 5

A simple encryption scheme known as Caesarian roll, is obtained by shifting characters a fixed amount. In Python we can obtain the order of a character with the `ord` function

    >>> ord('A')
    65
    
and there reverse operation

    >>> chr(65)
    'A' 
    
Use this in a function that takes a string and returns the encrypted representation

<!-- grade_id: cell-95b86ffef93462ec, points: 0, grade: false, solution: true -->
```python
def caesar_shift(text, n):
    """
    Returns an encrypted string by shifting input characters by n
    This is bloated but the basic idea is:
    - Split the string into list with length 1
    - Perform Caesar shift mod 26 (ASCII table)
    - Join the list
    """
    # YOUR CODE HERE
def caesar_shift(text, n):
    return ''.join(list(map(lambda x: chr(x), list(map(lambda x: (ord(x) + n - ((1 << 6) + (1 << 0))) % ((1 << 4) + (1 << 3) + (1 << 1)) + ((1 << 6) + (1 << 0)) if ((1 << 6) + (1 << 0)) <= ord(x) <= ((1 << 6) + (1 << 4) + (1 << 3) + (1 << 1) + (1 << 0)) else (ord(x)  + n - ((1 << 6) + (1 << 5) + (1 << 0))) % ((1 << 4) + (1 << 3) + (1 << 1)) + ((1 << 6) + (1 << 5) + (1 << 0)) if ((1 << 6) + (1 << 5) + (1 << 0)) <= ord(x) <= ((1 << 6) + (1 << 5) + (1 << 4) + (1 << 3) + (1 << 1) + (1 << 0)) else ord(x), [text[x:x + (1 << 0)] for x in range(0, len(text))])))))
    ################


print(caesar_shift('a', 0))
print(caesar_shift('ABBA', 2))
print(caesar_shift('CDDC', -2))
```

Output:
```
a
CDDC
ABBA
```

<!-- grade_id: cell-4403018cd26de943, points: 1, grade: true, solution: false -->
```python
assert caesar_shift('a', 0) == 'a'
assert caesar_shift('ABBA', 2) == 'CDDC'
```

<!-- grade_id: cell-d1495046a475f21c, points: 1, grade: true, solution: false -->
```python
assert caesar_shift('CDDC', -2) == 'ABBA'
```