Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Evan Huynh"
```

---

# File handling


## Creating our reference file

In our initial example we create a small reference file by executing the following cell. The so-called cell-magic command `%%file` takes its first argument as a filename and writes the remaining content of the cell to file

[More about cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html)

```python
%%file composers.txt
dufay
ockeghem
josquin

```

Output:
```
Overwriting composers.txt
```

## Reading

In the first exercise we recall different forms of reading text files. These are single-line answers that output some text. 



#### Exercise 1
What single line is used to return the file content as a single string? Not that the string will contain the end-of-line (newline) character `\n`

<!-- grade_id: cell-619ee4026f28565c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def a(filename):
    with open(filename) as f:
        return '\n'.join([line.rstrip('\n') for line in f.readlines()])
a('composers.txt')
################
```

Output:
```
'dufay\nockeghem\njosquin\n'
```

<!-- grade_id: cell-e16f081e726162ce, points: 1, grade: true, solution: false -->
```python
assert _ == '''\
dufay
ockeghem
josquin
'''
```

#### Exercise 2
What single line is used to return the file content as a Python list of strings. 

<!-- grade_id: cell-b8288e222d0343fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def a(filename):
    with open(filename) as f:
        return f.readlines()[:-1]
a('composers.txt')
################
```

Output:
```
['dufay\n', 'ockeghem\n', 'josquin\n']
```

<!-- grade_id: cell-3f4d59f2de6a1c5e, points: 1, grade: true, solution: false -->
```python
assert _ == ['dufay\n', 'ockeghem\n', 'josquin\n']
```

#### Exercise 3
This function will return the lines of the file without the 'end-of-line' characters. Use the [`strip`](https://docs.python.org/3/library/stdtypes.html#str.strip) method

<!-- grade_id: cell-a53c2938a84d15db, points: 0, grade: false, solution: true -->
```python
def get_list_of_lines_without_eol_character(filename):
    lines = []
    for line in open('composers.txt'):
        # YOUR CODE HERE
        lines.append(line.strip())
        ################
    lines = lines[:-1]
    return lines
 
get_list_of_lines_without_eol_character('composers.txt')
```

Output:
```
['dufay', 'ockeghem', 'josquin']
```

<!-- grade_id: cell-de85c7e2aeabc67f, points: 1, grade: true, solution: false -->
```python
assert get_list_of_lines_without_eol_character('composers.txt') == ['dufay', 'ockeghem', 'josquin']
```

## Reading numeric data
All reading/writing is done with text objects, i.e. reading a '7' will be text string with a single character rather than a numeric value. 

To get numbers from file intended for computation we have to do a type conversion (typically `int` or `float`)

Consider the following file with numbers. 


```python
%%file numbers1.txt
1
2
3
4
5
```

Output:
```
Overwriting numbers1.txt
```

### Exercise 4

Write a function that takes a file name as input and returns the sum of the numbers in the file, assuming one per line

<!-- grade_id: cell-e61d832d811f3cc4, points: 0, grade: false, solution: true -->
```python
def sum_numbers_in_file(filename):
    f = open(filename)
    number_sum = 0
    # YOUR CODE HERE
    for line in f:
        number_sum += int(line)
    ################
    f.close()
    return number_sum 
 
sum_numbers_in_file('numbers1.txt')
```

Output:
```
15
```

<!-- grade_id: cell-cedf19ef9aeb39f9, points: 1, grade: true, solution: false -->
```python

assert sum_numbers_in_file('numbers1.txt') == 15
```

### Exercise 5

It was realized that some files have more numbers on several lines such that for a file like

    1 2 3 4 5
    6 7 8 9 10
    
could not be handled by this function but we would get an error like

    E           ValueError: invalid literal for int() with base 10: '1 2 3 4 5\n'
    

In this exercise generalize the function above by copying here. Use the `split` method to split a line into words.

```python
%%file numbers2.txt
1 2 3 4 5
6 7 8 9 10


```

Output:
```
Overwriting numbers2.txt
```

```python
# You may uncomment this line and try
# sum_numbers_in_file('numbers2.txt')
```

In this exercise generalize the function from exercise 4. Use the `split` method to split a line into words.

<!-- grade_id: cell-e61d832d811f3cc5, points: 0, grade: false, solution: true -->
```python
def sum_numbers_in_file(filename):
    f = open(filename)
    number_sum = 0
    # 
    # YOUR CODE HERE
    for line in f:
        for i in line.split():
            number_sum += int(i)
    ################
    # COPY YOUR SOLUTION FROM EXERCISE 4 AND MODIFY
    f.close()
    return number_sum 
   
    
sum_numbers_in_file('numbers2.txt')
```

Output:
```
55
```

<!-- grade_id: cell-b4d65b9e6d0f1e01, points: 1, grade: true, solution: false -->
```python
assert sum_numbers_in_file('numbers1.txt') == 15
assert sum_numbers_in_file('numbers2.txt') == 55
```

## Writing files

When we call the builtin function `open` with a filename as a single argument the default behaviour is to look for an existing file for reading. If it does not exist the system prints  a `FileNotFoundError` and a message like

    open('no_such_file')
    ---------------------------------------------------------------------------
    FileNotFoundError                         Traceback (most recent call last)
    <ipython-input-15-b99a18bdaf59> in <module>()
    ----> 1 open('no_such_file')

    FileNotFoundError: [Errno 2] No such file or directory: 'no_such_file'
    
 
 
To open a new file to be written we have to supply an additional parameter, `w` for writing.

    open('new_file', 'w')
 
The `open` function returns a file object which has methods for reading (`read`, `readlines`) and writing (`write`)

You can often look up the documentation of a function with

```python
help(open)
```

Output:
```
Help on built-in function open in module io:

open(file, mode='r', buffering=-1, encoding=None, errors=None, newline=None, closefd=True, opener=None)
    Open file and return a stream.  Raise OSError upon failure.
    
    file is either a text or byte string giving the name (and the path
    if the file isn't in the current working directory) of the file to
    be opened or an integer file descriptor of the file to be
    wrapped. (If a file descriptor is given, it is closed when the
    returned I/O object is closed, unless closefd is set to False.)
    
    mode is an optional string that specifies the mode in which the file
    is opened. It defaults to 'r' which means open for reading in text
    mode.  Other common values are 'w' for writing (truncating the file if
    it already exists), 'x' for creating and writing to a new file, and
    'a' for appending (which on some Unix systems, means that all writes
    append to the end of the file regardless of the current seek position
```

### Exercise 6

With the content in the multi-line string write down code that writes the following content to a file of the form first_last.txt  and with capitalized names (First, Last) in the text.

<!-- grade_id: cell-9e5cf46774b3a243, points: 0, grade: false, solution: true -->
```python
def try_out_names(first, last):
    
    content = f"""
    'Tis but thy name that is my enemy;
    Thou art thyself, though not a Montague.
    What's {last}? It is nor hand, nor foot,
    Nor arm, nor face, nor any other part
    Belonging to a man. O, be some other name!
    What's in a name? That which we call a rose
    By any other word would smell as sweet;
    So {first} would, were he not {first} call'd,
    Retain that dear perfection which he owes
    Without that title. {first}, doff thy name,
    And for that name which is no part of thee
    Take all myself.
    """
    return content

def write_to_file(first, last):
    """
    Writes to file {first}_{last}.txt balcony scene with names
    """
    # YOUR CODE HERE
    with open(f'{first.lower()}_{last.lower()}.txt', 'w') as f:
        f.write(try_out_names(first.capitalize(), last.capitalize()))
    ################
    
    
print(try_out_names('Ronald', 'McDonald'))
write_to_file('Ronald', 'McDonald')
```

Output:
```
'Tis but thy name that is my enemy;
    Thou art thyself, though not a Montague.
    What's McDonald? It is nor hand, nor foot,
    Nor arm, nor face, nor any other part
    Belonging to a man. O, be some other name!
    What's in a name? That which we call a rose
    By any other word would smell as sweet;
    So Ronald would, were he not Ronald call'd,
    Retain that dear perfection which he owes
    Without that title. Ronald, doff thy name,
    And for that name which is no part of thee
    Take all myself.
```

<!-- grade_id: cell-afe447054344d72a, points: 1, grade: true, solution: false -->
```python
# This utility function removes an existing file if it exists
# like "rm -f"  (unix) or "del" (windows)
import pathlib
def remove_file_if_exists(filename):
    p = pathlib.Path(filename)
    if p.exists():
        p.unlink()


# Now write to new file and check contents

remove_file_if_exists('romeo_montague.txt')
write_to_file('ROMEO', 'MONTAGUE')

assert open('romeo_montague.txt').read() == try_out_names("Romeo", "Montague")
```

### Exercise 7

Writing numeric data to a file implies first converting to a string. The built-in function print does that. Look up the documentation on print how you can use it for writing to a file 

```python
help(print)
```

Output:
```
Help on built-in function print in module builtins:

print(...)
    print(value, ..., sep=' ', end='\n', file=sys.stdout, flush=False)
    
    Prints the values to a stream, or to sys.stdout by default.
    Optional keyword arguments:
    file:  a file-like object (stream); defaults to the current sys.stdout.
    sep:   string inserted between values, default a space.
    end:   string appended after the last value, default a newline.
    flush: whether to forcibly flush the stream.
```

<!-- grade_id: cell-df4516027c70782a, points: 0, grade: false, solution: true -->
```python
# Print the number pi to a file with pi.txt using print
from math import pi
# YOUR CODE HERE
print(pi, file=open('pi.txt', 'w'))
################
```

<!-- grade_id: cell-37e29477ff926375, points: 1, grade: true, solution: false -->
```python
# Check the first four characters in the file
assert open('pi.txt').read()[:4] == '3.14'
```

### Exercise 8
Implement a file copy command command with Python read an write commands


<!-- grade_id: cell-6360c6f2039531d7, points: 0, grade: false, solution: true -->
```python
def cp(source, target):
    """
    Copies textfile source to target
    """
    # YOUR CODE HERE
    with open(source, 'r') as f:
        with open(target, 'w') as g:
            g.write(f.read())
    ################
```

```python
%%file blake.txt
To see a World in a Grain of Sand
And a Heaven in a Wild Flower,
Hold Infinity in the palm of your hand 
And Eternity in an hour.
```

Output:
```
Overwriting blake.txt
```

<!-- grade_id: cell-d7bfacb6d05ced8a, points: 1, grade: true, solution: false -->
```python
cp('blake.txt', 'poem.txt')
assert open('blake.txt').read() == open('poem.txt').read()
```