Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
```

---

# Using classes in Python

* Date: 2022-05-06
* Deadline: 2022-05-20

### A. Create class for 3D vector 

Create a class for storing a 3D vector (x, y, z coordinates), and which provides a set of functions for manipulating the coordinates.  

* Implement class Vector3D which has the following properties:
* Attributes
    - coord_x stores x coordinate (object attribute)
    - coord_y stores y coordinate (object attribute)
    - coord_z stores z coordinate (object attribute)
* Ordinary class methods 
    - length() : computes length of vector* 
    - renormalize()   : renormalizes vector to unit length vector*
      * for square root use math.sqrt function from math library
* Overloaded operations  
    - addition operator(+)     : implement `__add__` method 
    - substraction operator(-) : implement `__sub__` method 

    

<!-- grade_id: cell-106104681ce80401, points: 0, grade: false, solution: true -->
```python
import math

class Vector3D:
    """
    Class for manipulation of 3D-vectors
    
    instance attributes:
        coord_x, coord_y, coord_z (float)
    
    class methods
        length (float): returns Euclidean norm of vector
        normalize (None): scales coordinates so that resulting vector has length 1
        __add__ (Vector3D) returns new vector as vector sum of two vectors
        __sub__ (Vector3D) returns new vector as vector difference of two vectors
        __repr__: string representation of vector

    """
    # YOUR CODE HERE
    ################

    
    
```

```python
# Define helper function
from pytest import approx
def error_message(actual, expected):
    return f'Actual {actual} != Expected {expected}'
```

<!-- grade_id: cell-e8becee65f92a575, points: 1, grade: true, solution: false -->
```python
# Verify Vector3D constructor
vector = Vector3D(1.0, 2.0, 4.0)

assert vector.coord_x == approx(1), error_message(vector.coord_x, 1)
assert vector.coord_y == approx(2), error_message(vector.coord_y, 2)
assert vector.coord_z == approx(4), error_message(vector.coord_z, 4)
```

<!-- grade_id: cell-3328f70c25343894, points: 1, grade: true, solution: false -->
```python
# Verify length method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

assert vector.length() == approx(math.sqrt(21)), error_message(vector.length(), math.sqrt(21))
```

<!-- grade_id: cell-96e5a932f3347dcf, points: 1, grade: true, solution: false -->
```python
# Verify norm method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

vector.renormalize()

assert vector.length() == approx(1), error_message(vector.length(), 1)
```

<!-- grade_id: cell-1f019ebe8894fb2b, points: 1, grade: true, solution: false -->
```python
#Verify addition operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 3.3, 0.9)

cvector = avector + bvector

assert cvector.coord_x == approx(3.0), error_message(cvector.coord_x, 3)
assert cvector.coord_y == approx(5.3), error_message(cvector.coord_y, 5.3)
assert cvector.coord_z == approx(4.9), error_message(cvector.coord_z, 4.9)

```

<!-- grade_id: cell-965fcb9cfdf9230b, points: 1, grade: true, solution: false -->
```python
#Verify substraction operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 0.2, 0.9)

cvector = avector - bvector

assert cvector.coord_x == approx(-1.0), error_message(cvector.coord_x, -1)
assert cvector.coord_y == approx(1.8), error_message(cvector.coord_y, 1.8)
assert cvector.coord_z == approx(3.1), error_message(cvector.coord_z, 3,1)
```

### B. Create class for currency data storage 

Create a class for storing money in SEK or Euros, which provides a set of functions for manipulating the amount and type of currency. 

* Implement class Money which has the following properties:
* Attributes
  - conversion_rate stores conversion rate from EUR to SEK (class attribute, default value 10.0)
  - currency stores currency type ('SEK' or 'EUR') (instance attribute)
  - amount stores amount of money (instance attribute)
* Ordinary class methods
  - set_conversion_rate(new_conversion_rate) : for setting conversion rate 
* Overloaded operations
  - addition operator(+)     : implement addition of money ammounts* 
  - substraction operator(-) : implement substraction of money ammounts* 
  * currency type of result must be currency type of first operand in arithmetical expresion 

<!-- grade_id: cell-dbe1d3027a05e5a4, points: 0, grade: false, solution: true -->
```python
class Money:
    """
    A class for manipulation of different currencies
    
    class attributes:
        conversion_rate (float) 
    instance attributes:
        amount (float)
        currency (str)
    class methods:
        set_conversion_rate: resets conversion rate
        __add__: implements addition of monies (right side currency converted to left side if different)
        __sub__: implements subtraction of monies
    """
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-73e7fae89242df32, points: 1, grade: true, solution: false -->
```python
#Verify class attribute
assert Money.conversion_rate == 10.0, error_message(Money.conversion_rate, 10)
```

<!-- grade_id: cell-4181fe5f998c9a96, points: 1, grade: true, solution: false -->
```python
#Verify instance attributes
ma = Money(100, 'SEK')
mb = Money(50, 'EUR')
assert ma.amount == 100, error_message(ma.amount, 100)
assert ma.currency == 'SEK', error_message(ma.currency, 'SEK')
assert mb.amount == 50, error_message(mb.amount, 50)
assert mb.currency == 'EUR', error_message(mb.currency, 'EUR')
```

<!-- grade_id: cell-fa09de3f68f4fad3, points: 1, grade: true, solution: false -->
```python
#Verify addition
ma = Money(100, 'SEK')
mb = Money(50, 'EUR')
mc = ma + mb
assert mc.amount == 600, error_message(mc.amount, 600)
assert mc.currency == 'SEK', error_message(mc.currency, 'SEK')
```

<!-- grade_id: cell-77d09f62ca89cb1d, points: 1, grade: true, solution: false -->
```python
#Verify subtraction
ma = Money(100, 'SEK')
mb = Money(50, 'EUR')
md = mb - ma
assert md.amount == 40, error_message(md.amount, 40)
assert md.currency == 'EUR', error_message(md.currency, 'EUR')
```

<!-- grade_id: cell-c2d199586e7c4cbb, points: 1, grade: true, solution: false -->
```python
#Set conversion rate
ma = Money(100, 'SEK')
mb = Money(50, 'EUR')
ma.set_conversion_rate(20)
assert ma.conversion_rate == 20, error_message(ma.conversion_rate, 20)
assert mb.conversion_rate == 20, error_message(mb.conversion_rate, 20)
```

## C. Animals

Complete the Animal class to have and a class variable `animals` (list) and two instance variables, `name` (str) and `number` (int). You need to implement `__init__` and `__str__` methods

<!-- grade_id: cell-f9faff46e9639b35, points: 0, grade: false, solution: true -->
```python
class Animal:
    """
    A class for storing animals
    
    class attributes:
        animals: (list) to store all animals
    instance attributes:
        name:  (str) to store animal name
        number: (int) to store animal order number (starting with 1)
        
    class methods:
        __str__: string representation of animal, e.g. "1. Dog"
        
    static methods:
        zoo: returns string representation of all animals in orderd lies, e.g. 
           '''\
           1. Dog
           2. Cat'''
    """
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-41580406aa20538d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

dog = Animal('dog')
assert dog.name == 'dog', error_message(dog.name, 'dog')
assert dog.number == 1, error_message(dog.number, 1)
assert str(dog) == '1. Dog', error_message(str(dog), '1. Dog')

cat = Animal('cat')
assert cat.name == 'cat', error_message(cat.name, 'cat')
assert cat.number == 2, error_message(cat.number, 2)
assert str(cat) == '2. Cat', error_message(str(cat), '2. Cat')

```

A static method is a function in a class that is like an ordinary function, that does not depend on any instance (no self argument). 

In a class it is defined with a `@staticmethod` decorator.

It can be appended to the class definition as below. Complete the static method so that it returns a string which lists all member animals

<!-- grade_id: cell-a89d339b99865f5d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

#Generate a list of animals and compare with the class attribute
animals = [Animal(a) for a in ['camel', 'donkey', 'hippo']]
assert animals == Animal.animals, f'{animals} != {Animal.animals}'

#zoo should produce a printout of the defined animals
zoo_output = Animal.zoo()
print(zoo_output)
expected_output = """
1. Camel
2. Donkey
3. Hippo
"""

condition = zoo_output.strip() == expected_output.strip()
error_message = f"\n{zoo_output}\n   !=\n{expected_output}"

assert condition, error_message
```